# Autograder Variables

The following code cell defines autograder-specific variables. Execute it and move on.

- INDUS proportion of non-retail business acres per town
- AGE proportion of owner-occupied units built prior to 1940

```python
import numpy as np
from keras.datasets import boston_housing
import numpy as np
np.random.seed(42)
import pandas as pd
from pandas.plotting import scatter_matrix
%matplotlib inline

def passed(): print('✅')

NB_TRAIN = 400
[X, y], _ = boston_housing.load_data()
X = (X - X.mean(axis=0)) / X.std(axis=0)
y = (y - y.mean()) / y.std()

X, y = X[:NB_TRAIN], y[:NB_TRAIN]
sorted_idxs = np.argsort(y)
X, y = X[sorted_idxs], y[sorted_idxs]
y = np.hstack([np.full(shape=100, fill_value=i) for i in range(4)])
idxs = np.arange(NB_TRAIN)
np.random.shuffle(idxs)
X, y = X[idxs], y[idxs]

X = X[:, [2, 6]]

NB_SAMPLE, NB_FEATURE = X.shape
classes = np.unique(y)
NB_CLASS = len(classes)
df = pd.DataFrame(X, columns=['INDUS', 'AGE'])
df['COST'] = y
scatter_matrix(df, diagonal='kde');
df[''] = ''
df[['INDUS', 'AGE', '', 'COST']]
```

# Task

- Define a class `MultiLayerPerceptron` which passes the Multi-Layer Perceptron tests

<!-- grade_id: mlp_solution, points: 0, grade: false, solution: true -->
```python
# YOUR CODE HERE
raise NotImplementedError()
```

# Constructor Tests

<!-- grade_id: constructor_tests, points: 1, grade: true, solution: false -->
```python
from classifiers import MultiLayerPerceptron

NB_HIDDEN = 16
mlp = MultiLayerPerceptron(NB_FEATURE, NB_HIDDEN, NB_CLASS)

assert type(mlp.W1) == np.ndarray
assert mlp.W1.shape == (NB_FEATURE, NB_HIDDEN)
assert mlp.W1.dtype == np.float64
assert type(mlp.b1) == np.ndarray
assert mlp.b1.shape == (NB_HIDDEN,)
assert mlp.b1.dtype == np.float64
assert type(mlp.W2) == np.ndarray
assert mlp.W2.shape == (NB_HIDDEN, NB_CLASS)
assert mlp.W2.dtype == np.float64
assert mlp.b2.shape == (NB_CLASS,)

passed()
```

# Prediction Test

<!-- grade_id: predict_tests, points: 1, grade: true, solution: false -->
```python
y_pred = mlp.predict(X)
assert y_pred.shape == (NB_SAMPLE, NB_CLASS)

passed()
```

# Evaluation Test

<!-- grade_id: evaluation_tests, points: 1, grade: true, solution: false -->
```python
acc = mlp.evaluate(X, y)
assert type(acc) == np.float64
assert 0 <= acc <= 1

P = mlp.predict(X)
y_pred = P.argmax(axis=1)
assert acc == np.mean(y_pred == y)

passed()
```

<!-- grade_id: random_description, points: 1, grade: true, solution: true -->
YOUR ANSWER HERE

<!-- grade_id: random_solution, points: 0, grade: false, solution: true -->
```python
# YOUR CODE HERE
raise NotImplementedError()
```

# Random Optimizer Test 

<!-- grade_id: random_tests, points: 1, grade: true, solution: false -->
```python
from classifiers import MultiLayerPerceptronWithRandomOptimizer

mlp = MultiLayerPerceptronWithRandomOptimizer(NB_FEATURE, NB_HIDDEN, NB_CLASS)

acc = mlp.evaluate(X, y)
for _ in range(1000):
    loss = mlp._forward(X, y)
    mlp.fit(X, y, nb_epoch=10)
    assert mlp._forward(X, y) <= loss
    
assert mlp.evaluate(X, y) > acc

passed()
```

# Task

- Define a class `MultiLayerPerceptronWithGradientDescentOptimizer` which extends your `classifiers.MultiLayerPerceptron` model with a gradient descent optimizer

# Hint

- Compute the gradient numerically with finite differences

<!-- grade_id: gd_solution, points: 0, grade: false, solution: true -->
```python
# YOUR CODE HERE
raise NotImplementedError()
```

# Gradient Descent Optimizer Test

<!-- grade_id: gd_tests, points: 1, grade: true, solution: false -->
```python
from classifiers import MultiLayerPerceptronWithGradientDescentOptimizer

mlp = MultiLayerPerceptronWithGradientDescentOptimizer(NB_FEATURE, NB_HIDDEN, NB_CLASS)

acc = mlp.evaluate(X, y)
loss = mlp._forward(X, y)
for _ in range(100):
    mlp.fit(X, y, nb_epoch=1)
    assert mlp._forward(X, y) < loss
    loss = mlp._forward(X, y)
    
assert mlp.evaluate(X, y) > acc

passed()
```